<br>
<br>
<div>
<img  style="float: left; padding-right: 100px; width: 350px" src="logai.png">
    </div>
    <br>
    <h2 align="center">LogAI</h2>
<hr>
<h4 align="center"><a href="https://t.co/IoL5C0LiPX">Zephania Reuben</a></h4>
<h4 align="center">September, 2025</h4>
<hr>
<h3 align="center">AfriAI Lab - ADVANCED AI SHORT COURSE TRAINING 2025
<br>
<hr>
DS9: TASK
</h3>
<hr>
<br>

<h3 align="center"> FINANCIAL INCLUSSION IN AFRICA - ZINDI COMPETITION </h3>

### Understand The Problem Statement

Financial Inclusion remains one of the main obstacles to economic and human development in Africa. For example, across Kenya, Rwanda, Tanzania, and Uganda only 9.1 million adults (or 13.9% of the adult population) have access to or use a commercial bank account.

Traditionally, access to bank accounts has been regarded as an indicator of financial inclusion. Despite the proliferation of mobile money in Africa, and the growth of innovative fintech solutions, banks still play a pivotal role in facilitating access to financial services. Access to bank accounts enable households to save and facilitate payments while also helping businesses build up their credit-worthiness and improve their access to other finance services. Therefore, access to bank accounts is an essential contributor to long-term economic growth.

The objective of this Dataset is to create a machine learning model to predict which individuals are most likely to have or use a bank account. The models and solutions developed can provide an indication of the state of financial inclusion in Kenya, Rwanda, Tanzania and Uganda, while providing insights into some of the key demographic factors that might drive individuals’ financial outcomes.

Data source available in the the zindi platform, [Zindi Africa](https://zindi.africa/competitions/financial-inclusion-in-africa)

### Hypothesis Generation
This is a very important stage in any data science/machine learning pipeline. It involves understanding the problem in detail by brainstorming as many factors as possible which can impact the outcome. It is done by understanding the problem statement thoroughly and before looking at the data.

Below are some of the factors which I think can affect the chance for a person to have a bank account

- People who have mobile phone have lower chance to use bank account because of mobile money services.
- People who are employed have a higher chance of having a bank account than People who are unemployed.
- People with low education have low chance to have bank account
- People in rural areas have low chance to have bank account 
- People who have age below 18 have low chance to have bank account
- Female have less chance to have bank account

```python
# Your work here!
```

Credit: [Davis David](https://twitter.com/Davis_McDavid)

Tests for q1 (points: 6)
```
>>> ypoly(3,3) == 139
True
```
```
>>> ypoly(-2,3) == 0
True
```


Tests for q2 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q3 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```


Tests for q4 (points: 5)
```
>>> assert_equal(longest_word("Hello there people"), 'people')
```


Tests for q5 (points: 5)
```
>>> assert_equal(dictionarize([1,2,2,3,4,5]), {1: 1, 2: 2, 3: 1, 4: 1, 5: 1})
```


Tests for q6 (points: 5)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```
